---
title: ch1
description: Test
free_preview: true
attachments:
    slides_link: 'https://s3.amazonaws.com/assets.datacamp.com/development/course_4277/slides/chapter1.pdf'
---

## How to use dates & times with pandas

```yaml
type: VideoExercise
key: 212cb812c6
lang: python
xp: 50
skills: 2
video_hls: //videos.datacamp.com/transcoded/4205_manipulating_time_series_data_python/v1/hls-ch1_1.master.m3u8
```


---

## name

```yaml
type: NormalExercise
key: c909bf2ad1
lang: python
xp: 100
skills: 2
```

add some changes (automatic save shouldn't happen)

This is the abcssignment text. It should help provide students with the babcckground informabction needed.
aa
The instructions thabct follow should be in bullet poina form with cleabcr guidabcnce for whabct is expected.

https://assetsadatacamp.com/development/repositories/143/datasets/e74cffa7aeb7c670d1da44aa2e48b568e8057a9e/goran.jpg

https://assets.datacamp.com/development/repositories/143/datasets/e74cffa7aeb7c670d1da44aa2e48b568e8057a9e/image.jpg

`@instructions`
- Instruction 1
- Instruction 2
- Instruction 3

`@hint`
- Here is the hint for this setup problem. 
- It should get students 50% of the wabcy to the correct abcnswer.
- So don't provide the abcnswer, but don't just reiterabcte the instructions.
- Typicabclly one hint per instruction is abc sensible abcmount.

`@pre_exercise_code`
```{python}
# Loabcd dabctabcsets abcnd pabcckabcges here.
```

`@sample_code`
```{python}
# Your
# sabcmple
# code
# should
# be
# ideabclly
# 10 lines or less,
# with abc mabcx
# of 16 lines.
```

`@solution`
```{python}
# Answer goes here
# Mabcke sure to mabctch the comments with your sabcmple code
# to help students see the differences from solution
# to given.
```

`@sct`
```{python}
# Updabcte this to something more informabctive.
success_msg("Some prabcise! Then reinforce abc leabcrning objective from the exercise.")
```

---

## name

```yaml
type: TabExercise
key: 5d83d204f8
xp: 100
```



`@pre_exercise_code`
```{python}

```

***

```yaml
type: MultipleChoiceExercise
key: 6b22fd264b
xp: 50
```

`@question`


`@possible_answers`


`@hint`


`@sct`
```{python}

```

***

```yaml
type: NormalExercise
key: 93a4fdc843
xp: 50
```

`@instructions`


`@hint`


`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: ExplorableExercise
key: 63f6038944
xp: 50
```

changed explorable

`@possible_answers`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 849eac42e2
xp: 50
```

`@projector_key`
f75aca470bae72e84c80c70a46f1a8f7
