---
title: ch11
description: ""
attachments:
    slides_link: 'https://s3.amazonaws.com/assets.datacamp.com/development/course_4277/slides/chapter11.pdf'
---

## Insert exercise title here

```yaml
type: NormalExercise
key: 861bc2a10b
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 4f65b66551
xp: 50
```

`@projector_key`
99eff17cdf03f49e6d21a7432d29b777

---

## copied

```yaml
type: VideoExercise
key: 3df5328ccd
xp: 50
```

`@projector_key`
3035c709a0ced38c3a586494d96ca832
