---
title: ch5
description: ""
attachments:
    slides_link: 'https://s3.amazonaws.com/assets.datacamp.com/development/course_4277/slides/chapter5.pdf'
---

## Insert exercise title here

```yaml
type: NormalExercise
key: 22d94f41e2
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: da9f30ccfe
xp: 50
```

`@projector_key`
7f3e3c7a1017e273d2189bbabd485635
