---
title: ch4
description: ""
attachments: {}
free_preview: true
---

## Insert exercise title here

```yaml
type: NormalExercise
key: a0c7291927
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
