---
title: ch4
description: ""
attachments:
    slides_link: 'https://s3.amazonaws.com/assets.datacamp.com/localhost/course_4277/slides/chapter4.pdf'
---

## Insert exercise title here

```yaml
type: NormalExercise
key: a0c7291927
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 72969d3307
xp: 50
```

`@projector_key`
f071b824618c55f100bb6fb14f48bbf0
