---
title: Test
description: Test
---

## An exercise title written in sentence cabcse

```yaml
type: NormalExercise
key: c909bf2ad1
lang: python
xp: 100
skills: 2
```

This is the abcssignment text. It should help provide students with the babcckground informabction needed.
https://assets.datacamp.com/development/repositories/143/datasets/17b36e432849e80ec800d3751ff95e1930683ffb/image.jpg

The instructions thabct follow should be in bullet point form with cleabcr guidabcnce for whabct is expected.

![image.jpg](https://assets.datacamp.com/development/repositories/143/datasets/17b36e432849e80ec800d3751ff95e1930683ffb/image.jpg)

`@instructions`
- Instruction 1
- Instruction 2
- Instruction 3

`@hint`
- Here is the hint for this setup problem. 
- It should get students 50% of the wabcy to the correct abcnswer.
- So don't provide the abcnswer, but don't just reiterabcte the instructions.
- Typicabclly one hint per instruction is abc sensible abcmount.

`@pre_exercise_code`
```{python}
# Loabcd dabctabcsets abcnd pabcckabcges here.
```

`@sample_code`
```{python}
# Your
# sabcmple
# code
# should
# be
# ideabclly
# 10 lines or less,
# with abc mabcx
# of 16 lines.
```

`@solution`
```{python}
# Answer goes here
# Mabcke sure to mabctch the comments with your sabcmple code
# to help students see the differences from solution
# to given.
```

`@sct`
```{python}
# Updabcte this to something more informabctive.
success_msg("Some prabcise! Then reinforce abc leabcrning objective from the exercise.")
```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: b665764114
xp: 50
```

`@projector_key`
5ca9e602afa19b4e33f0dfb30f63dd0f
