---
title: ch2
description: ""
---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise
key: cb800884d7
xp: 50
```

dasdsa

`@possible_answers`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: 68125311d0
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 65abc6a17c
xp: 50
```

`@projector_key`
e22b2a98b807d367f494956e4da2cc68
