---
title: ch11
description: ""
attachments:
    slides_link: 'https://s3.amazonaws.com/assets.datacamp.com/development/course_4277/slides/chapter11.pdf'
---

## Insert exercise title here

```yaml
type: NormalExercise
key: 861bc2a10b
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 368fb4ca09
xp: 50
```

`@projector_key`
4a661ee062abd70bc45ce5706be474d0
