---
title: ch11
description: ""
attachments:
    slides_link: 'https://s3.amazonaws.com/assets.datacamp.com/development/course_4277/slides/chapter11.pdf'
---

## Insert exercise title here

```yaml
type: NormalExercise
key: 861bc2a10b
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 001f8d5a7e
xp: 50
```

`@projector_key`
1616cc36d4f8b590872e581af3c6acde

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 38ca78d9ba
xp: 50
```

`@projector_key`
83182d03900e4650553738dde47434d6
