---
title: Test
description: Test
---

## An exercise title written in sentence cabcse

```yaml
type: NormalExercise
key: c909bf2ad1
lang: python
xp: 100
skills: 2
```

This is the abcssignment text. It should help provide students with the babcckground informabction needed.
aa
The instructions thabct follow should be in bullet point form with cleabcr guidabcnce for whabct is expected.

https://assets.datacamp.com/development/repositories/143/datasets/e74cffa7aeb7c670d1da44aa2e48b568e8057a9e/goran.jpg

https://assets.datacamp.com/development/repositories/143/datasets/e74cffa7aeb7c670d1da44aa2e48b568e8057a9e/image.jpg

`@instructions`
- Instruction 1
- Instruction 2
- Instruction 3

`@hint`
- Here is the hint for this setup problem. 
- It should get students 50% of the wabcy to the correct abcnswer.
- So don't provide the abcnswer, but don't just reiterabcte the instructions.
- Typicabclly one hint per instruction is abc sensible abcmount.

`@pre_exercise_code`
```{python}
# Loabcd dabctabcsets abcnd pabcckabcges here.
```

`@sample_code`
```{python}
# Your
# sabcmple
# code
# should
# be
# ideabclly
# 10 lines or less,
# with abc mabcx
# of 16 lines.
```

`@solution`
```{python}
# Answer goes here
# Mabcke sure to mabctch the comments with your sabcmple code
# to help students see the differences from solution
# to given.
```

`@sct`
```{python}
# Updabcte this to something more informabctive.
success_msg("Some prabcise! Then reinforce abc leabcrning objective from the exercise.")
```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: b665764114
xp: 50
```

`@projector_key`
5ca9e602afa19b4e33f0dfb30f63dd0f

---

## Insert exercise title here

```yaml
type: TabExercise
key: 5d83d204f8
xp: 100
```



`@pre_exercise_code`
```{python}

```

***

```yaml
type: MultipleChoiceExercise
key: 6b22fd264b
xp: 50
```

`@question`


`@possible_answers`


`@hint`


`@sct`
```{python}

```

***

```yaml
type: NormalExercise
key: 93a4fdc843
xp: 50
```

`@instructions`


`@hint`


`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## bla

```yaml
type: VideoExercise
key: 33ef76b6b7
xp: 50
```

`@projector_key`
3055c37537ac6e32a5a7fe9a1dd1c8e0
